<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 4763771  -->
  <question type="coderunner">
    <name>
      <text>Analyser un répertoire</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Écrire un script nommé <code>analyser</code> qui doit calculer le nombre de fichiers standard, de sous-répertoires, et d'exécutables d'un répertoire quelconque qui sera donné en paramètre.</p>
<p>Le répertoire ci-dessous (et d'autres caché) sera utiliser pour tester votre script :</p>
<pre>rep_exemple/
    fichier_std1 (fichier standard)
    fichier_std2 (fichier standard)
    fichier_exe  (fichier executable)
    sous_rep1/   (sous répertoire)
    sous_rep2/   (sous répertoire)
</pre>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<pre> </pre>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>bash</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>10, 20, ...</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload>#!/bin/bash</answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template><![CDATA[""" The template for a question type that compiles and runs a student-submitted
    bash script.
"""

import subprocess, sys, os


os.mkdir("./rep_exemple")
open("rep_exemple/fichier_std1", "w").close()
open("rep_exemple/fichier_std2", "w").close()
open("rep_exemple/fichier_exe", "w").close()
os.chmod("rep_exemple/fichier_exe", 0o755)
os.mkdir("./rep_exemple/sous_rep1")
os.mkdir("./rep_exemple/sous_rep2")

os.mkdir("./tmp")
open("tmp/f1", "w").close()
open("tmp/f2", "w").close()
open("tmp/f3", "w").close()
open("tmp/f4", "w").close()
os.chmod("tmp/f2", 0o755)
os.chmod("tmp/f3", 0o755)
os.mkdir("tmp/r1")
os.mkdir("tmp/r2")
os.mkdir("tmp/r3")

script_name="./analyser"

# Write the student code to a file script-etu.sh
student_answer = """{{ STUDENT_ANSWER | e('py') }}"""
with open(script_name, "w") as src:
    print(student_answer, file=src)

os.chmod(script_name, 0o755)

# Run the code. Since this is a per-test template,
# stdin is already set up for the stdin text specified in the test case,
# so we can run the compiled program directly.
try:
    args = [script_name]
    if len("{{ TEST.stdin  }}") :
        args.append("{{ TEST.stdin  }}")
    output = subprocess.check_output(args, universal_newlines=True)
    print(output)
except subprocess.CalledProcessError as e:
    if e.returncode > 0:
        # Ignore non-zero positive return codes
        if e.output:
            print(e.output)
    else:
        # But negative return codes are signals - abort
        if e.output:
            print(e.output, file=sys.stderr)
        if e.returncode < 0:
            print("Task failed with signal", -e.returncode, file=sys.stderr)
        print("** Further testing aborted **", file=sys.stderr)]]></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[#!/bin/bash
usage()
{
    echo "Usage: $0 <rep>"
    exit 1
}
[ "$#" -lt 1 ] && usage

rep=$1
c_dir=0
c_file=0
c_exe=0
for file in ${rep}/*
do
        if [ -d ${file} ]
        then
            echo "${file} est un sous-répertoire"
            c_dir=$((c_dir+1))
        fi
        if [ -f ${file} ]
        then
            if [ -x ${file} ]
            then
                echo "${file} est un fichier exécutable"
                c_exe=$((c_exe+1))
            else
                echo "${file} est un fichier standard"
                c_file=$((c_file+1))
            fi
        fi
done
echo "${rep}: dir: ${c_dir} file: ${c_file} file_exe: ${c_exe}"]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>./analyser</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text><![CDATA[Usage: ./analyser <rep>
]]></text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>./analyser</text>
      </testcode>
      <stdin>
                <text>rep_exemple</text>
      </stdin>
      <expected>
                <text>rep_exemple/fichier_exe est un fichier exécutable
rep_exemple/fichier_std1 est un fichier standard
rep_exemple/fichier_std2 est un fichier standard
rep_exemple/sous_rep1 est un sous-répertoire
rep_exemple/sous_rep2 est un sous-répertoire
rep_exemple: dir: 2 file: 2 file_exe: 1
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>./analyser</text>
      </testcode>
      <stdin>
                <text>tmp</text>
      </stdin>
      <expected>
                <text>tmp/f1 est un fichier standard
tmp/f2 est un fichier exécutable
tmp/f3 est un fichier exécutable
tmp/f4 est un fichier standard
tmp/r1 est un sous-répertoire
tmp/r2 est un sous-répertoire
tmp/r3 est un sous-répertoire
tmp: dir: 3 file: 2 file_exe: 2
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>HIDE</text>
      </display>
    </testcase>
    </testcases>
  </question>

</quiz>
